<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 2261545  -->
  <question type="coderunner">
    <name>
      <text>Faktorisierung</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>Schreiben Sie eine Funktion, die zu einer ganzen Zahl zwei mögliche Faktoren, die zusammen die übergebene Zahl ergeben, zurückgibt.<br>Verwenden Sie für die Funktion den folgenden Funktionskopf:<br></p><ul><li>void factorize(long product, long *f1, long *f2)<br></li></ul><p></p>Dabei soll folgender Algorithmus implementiert werden:<br><p></p><ul><li>product = 0 ---&gt; f1 = 1;&nbsp;&nbsp; f2 = product</li><li>product = 1 ---&gt; f1 = 1;&nbsp;&nbsp; f2 = product</li><li><strong>ansonsten</strong>: Wenn die Hälfte des Produktes keine ganzzahlige Vielfache vom Produkt selbst ist, verringere die Hälfte um 1 und <strong>prüfe mit der neuen Zahl erneut</strong><br>&nbsp;&nbsp;&nbsp;&nbsp; <strong>sobald </strong>es eine Vielfache des Produktes ist ---&gt; f1 = <em>die ganzzahlige Vielfache vom Produkt</em>;&nbsp;&nbsp; f2 = <em>das Produkt</em> / <em>die ganzzahlige Vielfache vom Produkt</em></li></ul><p><em><br></em></p><p>Da C-Funktionen immer nur <u>einen</u> Rückgabewert haben können (und wir 2 Werte zurückgeben wollen), soll die Wertrückgabe über Zeiger erfolgen.<br>Schreiben Sie außerdem eine <em>main</em>-Funktion, die eine ganze Zahl vom Benutzer einliest, diese mit Hilfe der Funktion in Faktoren zerlegt und anschließend die eingegebene Zahl und die Faktoren auf der Konsole ausgibt.<br><br><br></p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <coderunnertype>c_program</coderunnertype>
    <prototypetype>0</prototypetype>
    <allornothing>0</allornothing>
    <penaltyregime>0, 0, 5, 10, 20, ...</penaltyregime>
    <precheck>0</precheck>
    <hidecheck>0</hidecheck>
    <showsource>0</showsource>
    <answerboxlines>32</answerboxlines>
    <answerboxcolumns>100</answerboxcolumns>
    <answerpreload><![CDATA[#include <stdio.h>

void factorize(long product, long *f1, long *f2)
{
	// Your factorization
}

int main(){
	printf("Please enter an Integer: ");
	printf("Possible Factors of %ld are %ld and %ld.\n");
	return 0;
}
]]></answerpreload>
    <globalextra></globalextra>
    <useace></useace>
    <resultcolumns></resultcolumns>
    <template></template>
    <iscombinatortemplate></iscombinatortemplate>
    <allowmultiplestdins></allowmultiplestdins>
    <answer><![CDATA[#include <stdio.h>

void factorize(long product, long *f1, long *f2)
{
	unsigned long div = product/2;

	if(product == 0 || product == 1){
		*f1 = 1;
		*f2 = product;
		return;
	}
	while((product%div)!=0)
		div--;
	*f1 = div;
	*f2 = product/div;
	return;
}



int main(){
	long number = 0, factor1 = 0, factor2 = 0;
	printf("Please enter an Integer: ");
	scanf("%ld", &number);
	getchar();
	factorize(number, &factor1, &factor2);
	printf("Possible Factors of %ld are %ld and %ld.\n",
			number, factor1, factor2);
	return 0;
}
]]></answer>
    <validateonsave>1</validateonsave>
    <testsplitterre></testsplitterre>
    <language></language>
    <acelang></acelang>
    <sandbox></sandbox>
    <grader></grader>
    <cputimelimitsecs></cputimelimitsecs>
    <memlimitmb></memlimitmb>
    <sandboxparams><![CDATA[{ "compileargs": ["-std=c99"] }]]></sandboxparams>
    <templateparams></templateparams>
    <hoisttemplateparams>1</hoisttemplateparams>
    <templateparamslang>twig</templateparamslang>
    <templateparamsevalpertry>0</templateparamsevalpertry>
    <templateparamsevald></templateparamsevald>
    <twigall>0</twigall>
    <uiplugin></uiplugin>
    <uiparameters></uiparameters>
    <attachments>0</attachments>
    <attachmentsrequired>0</attachmentsrequired>
    <maxfilesize>10240</maxfilesize>
    <filenamesregex></filenamesregex>
    <filenamesexplain></filenamesexplain>
    <displayfeedback>1</displayfeedback>
    <giveupallowed>0</giveupallowed>
    <prototypeextra></prototypeextra>
    <testcases>
      <testcase testtype="0" useasexample="1" hiderestiffail="0" mark="1.0000000" >
      <testcode>
                <text>1. Test</text>
      </testcode>
      <stdin>
                <text>50</text>
      </stdin>
      <expected>
                <text>Please enter an Integer: Possible Factors of 50 are 25 and 2.
</text>
      </expected>
      <extra>
                <text></text>
      </extra>
      <display>
                <text>SHOW</text>
      </display>
    </testcase>
      <testcase testtype="0" useasexample="1" hiderestiffail="0" mark="1.0000000" >
      <testcode>
                <text>2. Test</text>
      </testcode>
      <stdin>
                <text>1</text>
      </stdin>
      <expected>
                <text>Please enter an Integer: Possible Factors of 1 are 1 and 1.
</text>
      </expected>
      <extra>
                <text></text>
      </extra>
      <display>
                <text>SHOW</text>
      </display>
    </testcase>
      <testcase testtype="0" useasexample="1" hiderestiffail="0" mark="1.0000000" >
      <testcode>
                <text>3. Test</text>
      </testcode>
      <stdin>
                <text>27</text>
      </stdin>
      <expected>
                <text>Please enter an Integer: Possible Factors of 27 are 9 and 3.
</text>
      </expected>
      <extra>
                <text></text>
      </extra>
      <display>
                <text>SHOW</text>
      </display>
    </testcase>
      <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000" >
      <testcode>
                <text>4. Test</text>
      </testcode>
      <stdin>
                <text>100</text>
      </stdin>
      <expected>
                <text>Please enter an Integer: Possible Factors of 100 are 50 and 2.
</text>
      </expected>
      <extra>
                <text></text>
      </extra>
      <display>
                <text>HIDE</text>
      </display>
    </testcase>
      <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000" >
      <testcode>
                <text>5. Test (Primzahl)</text>
      </testcode>
      <stdin>
                <text>19</text>
      </stdin>
      <expected>
                <text>Please enter an Integer: Possible Factors of 19 are 1 and 19.
</text>
      </expected>
      <extra>
                <text></text>
      </extra>
      <display>
                <text>HIDE_IF_SUCCEED</text>
      </display>
    </testcase>
    </testcases>
  </question>

</quiz>