<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 2168618  -->
  <question type="coderunner">
    <name>
      <text>Kontrollstrukturen</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>Schreiben Sie ein Programm, das eine Pizzabestellung simuliert. Nach einer Begrüssung soll dem Benutzer ein Menü praesentiert werden, in dem mindestens 3 Optionen zur Auswahl stehen, z.B. PizzaNr1 bestellen, PizzaNr2 bestellen und Bestellung abschliessen. Bei Auswahl jeder Option wird ein entsprechender Preis aufsummiert. Verwendet für PizzaNr1 den Betrag 5.50€ und für PizzaNr2 den Betrag 4.30€.<br>Nach Abschluss der Bestellung wird noch nach der Lieferentfernung gefragt, und ausgehend davon ein Rabatt gewährt, in Höhe von:<br></p><ul><li>10 Prozent, wenn die Entfernung maximal 2km beträgt, weil der Lieferbote dann laufen kann</li><li>10 Prozent, wenn die Entfernung mindestens 20km beträgt, weil die Pizza dann kalt ankommt</li><li><b>zusätzlichen</b> 10 Prozent, wenn die Bestellung einen Wert von 10 überschreitet</li><li>&nbsp;bzw. <b>zusätzlichen</b> 10 Prozent, wenn die Bestellung einen Wert von 20 überschreitet.<br></li></ul><p>Yum Yum Pizza braucht neue Kunden.<br><br>Der endgültige Preis wird am Ende präsentiert und der Nutzer verabschiedet.<br></p><p><span class="" style="color: rgb(255, 207, 53);">Hinweis</span>: Mit "%.2f" kann eine Float-Zahl mit nur ihren ersten 2 Nachkommastellen ausgegeben werden<br><br></p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text></text>
    </generalfeedback>
    <defaultgrade>3</defaultgrade>
    <penalty>0</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <coderunnertype>c_program</coderunnertype>
    <prototypetype>0</prototypetype>
    <allornothing>1</allornothing>
    <penaltyregime>0, 0, 0, 10, 20, ...</penaltyregime>
    <precheck>0</precheck>
    <hidecheck>0</hidecheck>
    <showsource>0</showsource>
    <answerboxlines>40</answerboxlines>
    <answerboxcolumns>100</answerboxcolumns>
    <answerpreload></answerpreload>
    <globalextra></globalextra>
    <useace></useace>
    <resultcolumns></resultcolumns>
    <template></template>
    <iscombinatortemplate></iscombinatortemplate>
    <allowmultiplestdins></allowmultiplestdins>
    <answer><![CDATA[#include <stdio.h>

int main()
{
  float gesamtpreis = 0;
  unsigned int entfernung = 0;
  int rabatt_prozent = 0;
  unsigned int auswahl = 0;
  printf("Ahoi bei Yum Yum Pizza!\n drücken Sie ...\n");
  while(auswahl != 3)
  {
    printf(" '1' fuer Salami Pizza dazu \n '2' fuer Veggi Pizza dazu \n '3' fuer Bestellung abschliessen \n");
    scanf( "%i", &auswahl );
    switch (auswahl)
    {
      case 1: 
        printf("Einmal Pizza Salami dazu, fuer 5.50\n");
        gesamtpreis += 5.50;
        break;
      case 2: 
        printf("Einmal Pizza Veggi dazu, fuer 4.30\n");
        gesamtpreis += 4.30;
        break;
    }
  }
  printf("\ndas wars? ...noch eine Frage: \nWelche Lieferentfernung hat die Bestellung (in km)? ");
  scanf( "%i", &entfernung );
  if(entfernung <= 2) rabatt_prozent += 10;
  if(entfernung >= 20) rabatt_prozent += 10;
  if (gesamtpreis > 10) rabatt_prozent += 10;
  if (gesamtpreis > 20) rabatt_prozent += 10;
  gesamtpreis -= gesamtpreis * (rabatt_prozent/100.0);
  printf("\nDas sind mit %i %% Rabatt genau %.2f Euro.\n", rabatt_prozent , gesamtpreis);
}
]]></answer>
    <validateonsave>1</validateonsave>
    <testsplitterre></testsplitterre>
    <language></language>
    <acelang></acelang>
    <sandbox></sandbox>
    <grader>NearEqualityGrader</grader>
    <cputimelimitsecs></cputimelimitsecs>
    <memlimitmb></memlimitmb>
    <sandboxparams></sandboxparams>
    <templateparams></templateparams>
    <hoisttemplateparams>1</hoisttemplateparams>
    <templateparamslang>twig</templateparamslang>
    <templateparamsevalpertry>0</templateparamsevalpertry>
    <templateparamsevald></templateparamsevald>
    <twigall>0</twigall>
    <uiplugin></uiplugin>
    <uiparameters></uiparameters>
    <attachments>0</attachments>
    <attachmentsrequired>0</attachmentsrequired>
    <maxfilesize>10240</maxfilesize>
    <filenamesregex></filenamesregex>
    <filenamesexplain></filenamesexplain>
    <displayfeedback>1</displayfeedback>
    <giveupallowed>0</giveupallowed>
    <prototypeextra></prototypeextra>
    <testcases>
      <testcase testtype="0" useasexample="1" hiderestiffail="0" mark="1.0000000" >
      <testcode>
                <text>1. Kunde (mittelgroße Bestellung, kurze Entfernung)</text>
      </testcode>
      <stdin>
                <text>1
1
3
1
</text>
      </stdin>
      <expected>
                <text>Ahoi bei Yum Yum Pizza!
 drücken Sie ...
 '1' fuer Salami Pizza dazu
 '2' fuer Veggi Pizza dazu
 '3' fuer Bestellung abschliessen
Einmal Pizza Salami dazu, fuer 5.50
 '1' fuer Salami Pizza dazu
 '2' fuer Veggi Pizza dazu
 '3' fuer Bestellung abschliessen
Einmal Pizza Salami dazu, fuer 5.50
 '1' fuer Salami Pizza dazu
 '2' fuer Veggi Pizza dazu
 '3' fuer Bestellung abschliessen

das wars? ...noch eine Frage:
Welche Lieferentfernung hat die Bestellung (in km)?
Das sind mit 20 % Rabatt genau 8.80 Euro.
</text>
      </expected>
      <extra>
                <text></text>
      </extra>
      <display>
                <text>SHOW</text>
      </display>
    </testcase>
      <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000" >
      <testcode>
                <text>2. Kunde (kleine Bestellung, kurze Entfernung)</text>
      </testcode>
      <stdin>
                <text>1
3
2
</text>
      </stdin>
      <expected>
                <text>Ahoi bei Yum Yum Pizza!
 drücken Sie ...
 '1' fuer Salami Pizza dazu
 '2' fuer Veggi Pizza dazu
 '3' fuer Bestellung abschliessen
Einmal Pizza Salami dazu, fuer 5.50
 '1' fuer Salami Pizza dazu
 '2' fuer Veggi Pizza dazu
 '3' fuer Bestellung abschliessen

das wars? ...noch eine Frage:
Welche Lieferentfernung hat die Bestellung (in km)?
Das sind mit 10 % Rabatt genau 4.95 Euro.
</text>
      </expected>
      <extra>
                <text></text>
      </extra>
      <display>
                <text>HIDE_IF_SUCCEED</text>
      </display>
    </testcase>
      <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000" >
      <testcode>
                <text>3. Kunde (kleine Bestellung, normale Entfernung)</text>
      </testcode>
      <stdin>
                <text>2
3
19
</text>
      </stdin>
      <expected>
                <text>Ahoi bei Yum Yum Pizza!
 drücken Sie ...
 '1' fuer Salami Pizza dazu
 '2' fuer Veggi Pizza dazu
 '3' fuer Bestellung abschliessen
Einmal Pizza Veggi dazu, fuer 4.30
 '1' fuer Salami Pizza dazu
 '2' fuer Veggi Pizza dazu
 '3' fuer Bestellung abschliessen

das wars? ...noch eine Frage:
Welche Lieferentfernung hat die Bestellung (in km)?
Das sind mit 0 % Rabatt genau 4.30 Euro.
</text>
      </expected>
      <extra>
                <text></text>
      </extra>
      <display>
                <text>HIDE_IF_SUCCEED</text>
      </display>
    </testcase>
      <testcase testtype="0" useasexample="1" hiderestiffail="0" mark="1.0000000" >
      <testcode>
                <text>4. Kunde (große Bestellung, große Entfernung)</text>
      </testcode>
      <stdin>
                <text>1
2
1
1
1
3
20
</text>
      </stdin>
      <expected>
                <text>Ahoi bei Yum Yum Pizza!
 drücken Sie ...
 '1' fuer Salami Pizza dazu
 '2' fuer Veggi Pizza dazu
 '3' fuer Bestellung abschliessen
Einmal Pizza Salami dazu, fuer 5.50
 '1' fuer Salami Pizza dazu
 '2' fuer Veggi Pizza dazu
 '3' fuer Bestellung abschliessen
Einmal Pizza Veggi dazu, fuer 4.30
 '1' fuer Salami Pizza dazu
 '2' fuer Veggi Pizza dazu
 '3' fuer Bestellung abschliessen
Einmal Pizza Salami dazu, fuer 5.50
 '1' fuer Salami Pizza dazu
 '2' fuer Veggi Pizza dazu
 '3' fuer Bestellung abschliessen
Einmal Pizza Salami dazu, fuer 5.50
 '1' fuer Salami Pizza dazu
 '2' fuer Veggi Pizza dazu
 '3' fuer Bestellung abschliessen
Einmal Pizza Salami dazu, fuer 5.50
 '1' fuer Salami Pizza dazu
 '2' fuer Veggi Pizza dazu
 '3' fuer Bestellung abschliessen

das wars? ...noch eine Frage:
Welche Lieferentfernung hat die Bestellung (in km)?
Das sind mit 30 % Rabatt genau 18.41 Euro.
</text>
      </expected>
      <extra>
                <text></text>
      </extra>
      <display>
                <text>HIDE_IF_SUCCEED</text>
      </display>
    </testcase>
      <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000" >
      <testcode>
                <text>5. Kunde (kleine Bestellung, große Entfernung)</text>
      </testcode>
      <stdin>
                <text>1
3
200
</text>
      </stdin>
      <expected>
                <text>Ahoi bei Yum Yum Pizza!
 drücken Sie ...
 '1' fuer Salami Pizza dazu
 '2' fuer Veggi Pizza dazu
 '3' fuer Bestellung abschliessen
Einmal Pizza Salami dazu, fuer 5.50
 '1' fuer Salami Pizza dazu
 '2' fuer Veggi Pizza dazu
 '3' fuer Bestellung abschliessen

das wars? ...noch eine Frage:
Welche Lieferentfernung hat die Bestellung (in km)?
Das sind mit 10 % Rabatt genau 4.95 Euro.
</text>
      </expected>
      <extra>
                <text></text>
      </extra>
      <display>
                <text>HIDE_IF_SUCCEED</text>
      </display>
    </testcase>
    </testcases>
  </question>

</quiz>