<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 2261555  -->
  <question type="coderunner">
    <name>
      <text>Struktur Student</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>Deklariert zunächst eine Struktur <b>struct student</b> und damit den Typ <b>Student</b> (typedef). Diese soll folgende Elemente enthalten (mit folgenden <em>Variablennamen</em>):<br></p><ul><li>Einen Vornamen von maximal 20 Zeichen Länge (<em>vorname</em>)<br></li><li>Einen Nachnamen von maximal 20 Zeichen Länge (<em>nachname</em>)<br></li><li>Einen Array, in welchem die Punkte von 5 Hausaufgabenblättern gespeichert werden können (Nur ganzzahlige Punkte sind erforderlich) (<em>punkte</em>)<br></li></ul><p>Schreibt eine <b>main()</b>-Funktion, welche zuerst ein Array mit Platz für 5 Studierende erstellt. Ruft nun für jeden Studenten eine (selbst zu schreibende!) Funktion <b><em>void student_einlesen(Student* p_stud)</em></b> auf, welche die Daten vom Nutzer einliest, in die Struktur schreibt und diese dann in dem Array speichert. Wenn alle Daten eingelesen wurden, ruft für jeden Studenten die (selbst zu schreibende...) Funktion <em><b>void student_ausgeben(</b><b>Student* p_stud</b><b>)</b></em> auf, welche alle Daten der übergebenen Struktur im Terminal ausgibt. Die Übergabe der Struktur soll by-reference, also als Zeiger erfolgen. Anstatt der Punkte jedes Aufgabenbogens soll die Summe aller Punkte ausgegeben werden.<br></p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text></text>
    </generalfeedback>
    <defaultgrade>7</defaultgrade>
    <penalty>0</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <coderunnertype>c_program</coderunnertype>
    <prototypetype>0</prototypetype>
    <allornothing>0</allornothing>
    <penaltyregime>0, 0, 5, 10, 20, ...</penaltyregime>
    <precheck>0</precheck>
    <hidecheck>0</hidecheck>
    <showsource>0</showsource>
    <answerboxlines>40</answerboxlines>
    <answerboxcolumns>100</answerboxcolumns>
    <answerpreload></answerpreload>
    <globalextra></globalextra>
    <useace></useace>
    <resultcolumns></resultcolumns>
    <template><![CDATA[{# name of the function you want to parse (flexibel) #}
{% set function_name = "main" %}
{# parse STUDENT_ANSWER #}
{% set text_after_function_name = STUDENT_ANSWER | split(function_name)[1] %}
{# constants (DO NOT CHANGE) #}
{% set occurence_of_open = 0 %}
{% set occurence_of_close = 0 %}
{% set function_string = "" %}
{% set parsing_done = false %}
{# outer loop { #}
{% for text_till_next_opened_brace in text_after_function_name | split("{")  %}
    {# Count how many { HAVE ALREADY occured.#}
    {% if not loop.first %}
        {% set occurence_of_open = occurence_of_open + 1 %}
    {% endif %}
    {# add { to string. First not because it has occured before first { second not, as we only want the content inside the {} #}
    {% if not parsing_done and loop.index0 >=2%}
        {% set function_string = function_string ~ "{" %}
    {% endif %}
    {# we need the information later #}
    {% set outer_loop_first = loop.first %}
    {# inner loop } #}
    {% for text_till_next_closed_brace in text_till_next_opened_brace | split("}")  %}
        {# Count how many } HAVE ALREADY occured.#}
        {% if not loop.first %}
            {% set occurence_of_close = occurence_of_close + 1 %}
        {% endif %}
        {# if this criterium holds, function is over } #}
        {% if occurence_of_open <= occurence_of_close and occurence_of_open > 0 %}
            {% set parsing_done = true %}
        {% endif %}
        {# content before and after fucntion is ignored } #}
        {% if not parsing_done and not outer_loop_first%}
            {# add } to string  #}
            {% if not loop.first %}
                {% set function_string = function_string ~ "}" %}
            {% endif %}
            {# add CONTENT to the string #}
            {% set function_string = function_string ~ text_till_next_closed_brace %}
        {% endif %}
    {% endfor %}
{% endfor %}

{% if "return" in function_string %}
    {% set function_has_return = true %}
{% else %}
    {% set function_has_return = false %}
{% endif %}

{% if function_has_return %}
{{ STUDENT_ANSWER | replace({'int main': 'int student_main'}) }} 
{% else %}
{{ STUDENT_ANSWER | replace({'int main': 'void student_main', 'void main': 'void student_main'}) }} 
{% endif %}


int main() {
    {{TEST.testcode}}
}

]]></template>
    <iscombinatortemplate></iscombinatortemplate>
    <allowmultiplestdins></allowmultiplestdins>
    <answer><![CDATA[#include <stdio.h>

typedef struct student{
    char vorname[21];
    char nachname[21];
    int punkte[5];
} Student;

void student_einlesen(Student* p_stud)
{
    printf("\nGebe den Vornamen ein: ");
    scanf("%s", p_stud->vorname);
    printf("Gebe nun den Nachnamen ein: ");
    scanf("%s", p_stud->nachname);
    printf("Gebe nun nacheinander die 5 Punkte ein:\n");
    for(int i=0; i<5; i++)
    {
        printf("%i: ", i+1);
        scanf("%i", &(p_stud->punkte[i]));
    }
}

void student_ausgeben(Student* p_stud)
{
    int sum = 0;
    for(int i=0; i<5; i++)
        sum += p_stud->punkte[i];
    printf("\nStudent: %s %s\n", p_stud->vorname, p_stud->nachname );
    printf("Anzahl Punkte: %i\n", sum);
}

int main()
{
    Student st[5];
    for (int i=0; i<5; i++)
        student_einlesen(&st[i]);
    for (int i=0; i<5; i++)
        student_ausgeben(&st[i]);
    return 0;
}
]]></answer>
    <validateonsave>1</validateonsave>
    <testsplitterre></testsplitterre>
    <language></language>
    <acelang></acelang>
    <sandbox></sandbox>
    <grader>NearEqualityGrader</grader>
    <cputimelimitsecs></cputimelimitsecs>
    <memlimitmb></memlimitmb>
    <sandboxparams><![CDATA[{ "compileargs": ["-std=c99"] }]]></sandboxparams>
    <templateparams></templateparams>
    <hoisttemplateparams>1</hoisttemplateparams>
    <templateparamslang>twig</templateparamslang>
    <templateparamsevalpertry>0</templateparamsevalpertry>
    <templateparamsevald></templateparamsevald>
    <twigall>0</twigall>
    <uiplugin></uiplugin>
    <uiparameters></uiparameters>
    <attachments>0</attachments>
    <attachmentsrequired>0</attachmentsrequired>
    <maxfilesize>10240</maxfilesize>
    <filenamesregex></filenamesregex>
    <filenamesexplain></filenamesexplain>
    <displayfeedback>1</displayfeedback>
    <giveupallowed>0</giveupallowed>
    <prototypeextra></prototypeextra>
    <testcases>
      <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000" >
      <testcode>
                <text><![CDATA[Student erster_Student;
student_einlesen(&erster_Student);
printf("\nerster_Student: %s %s\n", erster_Student.vorname, erster_Student.nachname);
]]></text>
      </testcode>
      <stdin>
                <text>Marco
Wagner
14
18
13
9
16</text>
      </stdin>
      <expected>
                <text>Gebe den Vornamen ein: Gebe nun den Nachnamen ein: Gebe nun nacheinander die 5 Punkte ein:
1: 2: 3: 4: 5:
erster_Student: Marco Wagner
</text>
      </expected>
      <extra>
                <text></text>
      </extra>
      <display>
                <text>SHOW</text>
      </display>
    </testcase>
      <testcase testtype="0" useasexample="1" hiderestiffail="0" mark="1.0000000" >
      <testcode>
                <text>student_main();</text>
      </testcode>
      <stdin>
                <text>Spongebob
Schwammkopf
19
20
18
20
20
Patrick
Star
5
6
3
20
5
Sandy
Cheeks
14
17
18
18
17
Thaddäus
Tentakel
10
12
11
14
10
Eugene
Krabs
8
4
7
10
13
</text>
      </stdin>
      <expected>
                <text>Gebe den Vornamen ein: Gebe nun den Nachnamen ein: Gebe nun nacheinander die 5 Punkte ein:
1: 2: 3: 4: 5:
Gebe den Vornamen ein: Gebe nun den Nachnamen ein: Gebe nun nacheinander die 5 Punkte ein:
1: 2: 3: 4: 5:
Gebe den Vornamen ein: Gebe nun den Nachnamen ein: Gebe nun nacheinander die 5 Punkte ein:
1: 2: 3: 4: 5:
Gebe den Vornamen ein: Gebe nun den Nachnamen ein: Gebe nun nacheinander die 5 Punkte ein:
1: 2: 3: 4: 5:
Gebe den Vornamen ein: Gebe nun den Nachnamen ein: Gebe nun nacheinander die 5 Punkte ein:
1: 2: 3: 4: 5:
Student: Spongebob Schwammkopf
Anzahl Punkte: 97

Student: Patrick Star
Anzahl Punkte: 39

Student: Sandy Cheeks
Anzahl Punkte: 84

Student: Thaddäus Tentakel
Anzahl Punkte: 57

Student: Eugene Krabs
Anzahl Punkte: 42
</text>
      </expected>
      <extra>
                <text></text>
      </extra>
      <display>
                <text>SHOW</text>
      </display>
    </testcase>
      <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000" >
      <testcode>
                <text>student_main();</text>
      </testcode>
      <stdin>
                <text>John
Snow
20
14
16
17
19
Robb
Stark
17
18
19
20
20
Arya 
Stark
3
5
9
12
10
Bran
Stark
14
17
13
17
17
Rickon
Stark
5
4
3
7
9
</text>
      </stdin>
      <expected>
                <text>Gebe den Vornamen ein: Gebe nun den Nachnamen ein: Gebe nun nacheinander die 5 Punkte ein:
1: 2: 3: 4: 5:
Gebe den Vornamen ein: Gebe nun den Nachnamen ein: Gebe nun nacheinander die 5 Punkte ein:
1: 2: 3: 4: 5:
Gebe den Vornamen ein: Gebe nun den Nachnamen ein: Gebe nun nacheinander die 5 Punkte ein:
1: 2: 3: 4: 5:
Gebe den Vornamen ein: Gebe nun den Nachnamen ein: Gebe nun nacheinander die 5 Punkte ein:
1: 2: 3: 4: 5:
Gebe den Vornamen ein: Gebe nun den Nachnamen ein: Gebe nun nacheinander die 5 Punkte ein:
1: 2: 3: 4: 5:
Student: John Snow
Anzahl Punkte: 86

Student: Robb Stark
Anzahl Punkte: 94

Student: Arya Stark
Anzahl Punkte: 39

Student: Bran Stark
Anzahl Punkte: 78

Student: Rickon Stark
Anzahl Punkte: 28
</text>
      </expected>
      <extra>
                <text></text>
      </extra>
      <display>
                <text>HIDE_IF_SUCCEED</text>
      </display>
    </testcase>
    </testcases>
  </question>

</quiz>