<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 2159942  -->
  <question type="coderunner">
    <name>
      <text>Schaltjahr</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>Schreibt ein Programm, dass feststellen kann, ob eine eingegebene Zahl einem Schaltjahr entspricht. Ein Jahr sei ein Schaltjahr, wenn die Jahreszahl durch 400 teilbar ist, oder wenn es durch 4, aber nicht durch 100 teilbar ist.<br></p><p>Die <span class="" style="color: rgb(126, 255, 126);"><span class="" style="color: rgb(125, 159, 211);">Ausgaben</span> </span>sollen <b>exakt </b>wie in den Beispieltests ausgegeben werden. Also:</p><p>"<span class="" style="color: rgb(125, 159, 211);">Das Programm kann Schaltjahre bestimmen - welches Jahr soll getestet werden?</span>"</p><p>Und eine weitere <span class="" style="color: rgb(126, 255, 126);"><span class="" style="color: rgb(125, 159, 211);">Ausgabe</span> </span>mit der <span class="" style="color: rgb(247, 163, 161);">Eingabe </span>vom Nutzer und dem anschließend Erebnis mit<br>"<span class="" style="color: rgb(126, 255, 126);"><span class="" style="color: rgb(247, 163, 161);">[Zahl]</span> <span class="" style="color: rgb(125, 159, 211);">ist ein</span><span class="" style="color: rgb(125, 159, 211);">/</span><span class="" style="color: rgb(125, 159, 211);">kein Schaltjahr</span>.</span>"<br><span><span class="" style="color: rgb(255, 207, 53);">Hinweis</span><span>: 0 &lt; eingegebene Jahr</span></span><br></p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text></text>
    </generalfeedback>
    <defaultgrade>2</defaultgrade>
    <penalty>0</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <coderunnertype>c_program</coderunnertype>
    <prototypetype>0</prototypetype>
    <allornothing>1</allornothing>
    <penaltyregime>0, 10, 20, ...</penaltyregime>
    <precheck>0</precheck>
    <hidecheck>0</hidecheck>
    <showsource>0</showsource>
    <answerboxlines>18</answerboxlines>
    <answerboxcolumns>100</answerboxcolumns>
    <answerpreload></answerpreload>
    <globalextra></globalextra>
    <useace></useace>
    <resultcolumns></resultcolumns>
    <template></template>
    <iscombinatortemplate></iscombinatortemplate>
    <allowmultiplestdins></allowmultiplestdins>
    <answer><![CDATA[#include <stdio.h>
 
int main()
{
    unsigned int jahreszahl;
   
    printf("Das Programm kann Schaltjahre bestimmen - welches Jahr soll getestet werden? ");
    scanf("%i", &jahreszahl);
   
    if (jahreszahl % 400 == 0) printf("\n%i ist ein Schaltjahr.\n", jahreszahl); //durch 400 ohne Rest teilbar?
    else if ((jahreszahl % 100 != 0) && (jahreszahl % 4 == 0))
        printf("\n%i ist ein Schaltjahr.\n", jahreszahl); //alle 4
    else printf("\n%i ist kein Schaltjahr.\n", jahreszahl); // alles andere
}]]></answer>
    <validateonsave>1</validateonsave>
    <testsplitterre></testsplitterre>
    <language></language>
    <acelang></acelang>
    <sandbox></sandbox>
    <grader></grader>
    <cputimelimitsecs></cputimelimitsecs>
    <memlimitmb></memlimitmb>
    <sandboxparams></sandboxparams>
    <templateparams></templateparams>
    <hoisttemplateparams>1</hoisttemplateparams>
    <templateparamslang>twig</templateparamslang>
    <templateparamsevalpertry>0</templateparamsevalpertry>
    <templateparamsevald></templateparamsevald>
    <twigall>0</twigall>
    <uiplugin></uiplugin>
    <uiparameters></uiparameters>
    <attachments>0</attachments>
    <attachmentsrequired>0</attachmentsrequired>
    <maxfilesize>10240</maxfilesize>
    <filenamesregex></filenamesregex>
    <filenamesexplain></filenamesexplain>
    <displayfeedback>1</displayfeedback>
    <giveupallowed>0</giveupallowed>
    <prototypeextra></prototypeextra>
    <testcases>
      <testcase testtype="0" useasexample="1" hiderestiffail="0" mark="1.0000000" >
      <testcode>
                <text>Ist 1900 ein Schaltjahr?</text>
      </testcode>
      <stdin>
                <text>1900</text>
      </stdin>
      <expected>
                <text>Das Programm kann Schaltjahre bestimmen - welches Jahr soll getestet werden?
1900 ist kein Schaltjahr.
</text>
      </expected>
      <extra>
                <text></text>
      </extra>
      <display>
                <text>SHOW</text>
      </display>
    </testcase>
      <testcase testtype="0" useasexample="1" hiderestiffail="0" mark="1.0000000" >
      <testcode>
                <text>Ist 2000 ein Schaltjahr?</text>
      </testcode>
      <stdin>
                <text>2000</text>
      </stdin>
      <expected>
                <text>Das Programm kann Schaltjahre bestimmen - welches Jahr soll getestet werden?
2000 ist ein Schaltjahr.
</text>
      </expected>
      <extra>
                <text></text>
      </extra>
      <display>
                <text>SHOW</text>
      </display>
    </testcase>
      <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000" >
      <testcode>
                <text>Ist 1 ein Schaltjahr?</text>
      </testcode>
      <stdin>
                <text>1</text>
      </stdin>
      <expected>
                <text>Das Programm kann Schaltjahre bestimmen - welches Jahr soll getestet werden?
1 ist kein Schaltjahr.
</text>
      </expected>
      <extra>
                <text></text>
      </extra>
      <display>
                <text>HIDE_IF_SUCCEED</text>
      </display>
    </testcase>
      <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000" >
      <testcode>
                <text>Ist 2004 ein Schaltjahr?</text>
      </testcode>
      <stdin>
                <text>2004</text>
      </stdin>
      <expected>
                <text>Das Programm kann Schaltjahre bestimmen - welches Jahr soll getestet werden?
2004 ist ein Schaltjahr.
</text>
      </expected>
      <extra>
                <text></text>
      </extra>
      <display>
                <text>HIDE_IF_SUCCEED</text>
      </display>
    </testcase>
      <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000" >
      <testcode>
                <text>Ist 2005 ein Schaltjahr?</text>
      </testcode>
      <stdin>
                <text>2005</text>
      </stdin>
      <expected>
                <text>Das Programm kann Schaltjahre bestimmen - welches Jahr soll getestet werden?
2005 ist kein Schaltjahr.
</text>
      </expected>
      <extra>
                <text></text>
      </extra>
      <display>
                <text>HIDE_IF_SUCCEED</text>
      </display>
    </testcase>
    </testcases>
  </question>

</quiz>