<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 3420296  -->
  <question type="coderunner">
    <name>
      <text>Ersetze while/if (Neu, FERTIG)</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>Ersetze <b>if</b>, <b>do...while</b> durch eine <b>for</b>-Konstruktion.<br></p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <coderunnertype>c_advanced</coderunnertype>
    <prototypetype>0</prototypetype>
    <allornothing>0</allornothing>
    <penaltyregime>0, 0, 5, 10, 20, ...</penaltyregime>
    <precheck>0</precheck>
    <hidecheck>0</hidecheck>
    <showsource>0</showsource>
    <answerboxlines>18</answerboxlines>
    <answerboxcolumns>100</answerboxcolumns>
    <answerpreload><![CDATA[#include <stdio.h>
#include <unistd.h>

int main(){
    printf( "Wieviele Sekunden soll ich warten?\n");
    int zeit=0;
    scanf("%i",&zeit);
    if (zeit >= 0){
        do{
        //eine Sekunde warten
        sleep(1);
        printf( "time %i\n", zeit);
        zeit--;
        } while (zeit >= 0);
    }
}]]></answerpreload>
    <globalextra></globalextra>
    <useace></useace>
    <resultcolumns></resultcolumns>
    <template></template>
    <iscombinatortemplate></iscombinatortemplate>
    <allowmultiplestdins></allowmultiplestdins>
    <answer><![CDATA[#include <stdio.h>
#include <unistd.h>

int main()
{
    printf( "Wieviele Sekunden soll ich warten?\n");
    int zeit=0;
    scanf("%i",&zeit);
    for (;zeit >= 0; zeit--){
        //eine Sekunde warten
        sleep(1);
        printf( "time %i\n", zeit);
    }
}]]></answer>
    <validateonsave>1</validateonsave>
    <testsplitterre></testsplitterre>
    <language></language>
    <acelang></acelang>
    <sandbox></sandbox>
    <grader></grader>
    <cputimelimitsecs></cputimelimitsecs>
    <memlimitmb></memlimitmb>
    <sandboxparams></sandboxparams>
    <templateparams></templateparams>
    <hoisttemplateparams>1</hoisttemplateparams>
    <templateparamslang>twig</templateparamslang>
    <templateparamsevalpertry>0</templateparamsevalpertry>
    <templateparamsevald>{}</templateparamsevald>
    <twigall>0</twigall>
    <uiplugin></uiplugin>
    <uiparameters></uiparameters>
    <attachments>0</attachments>
    <attachmentsrequired>0</attachmentsrequired>
    <maxfilesize>10240</maxfilesize>
    <filenamesregex></filenamesregex>
    <filenamesexplain></filenamesexplain>
    <displayfeedback>1</displayfeedback>
    <giveupallowed>0</giveupallowed>
    <prototypeextra></prototypeextra>
    <testcases>
      <testcase testtype="0" useasexample="0" hiderestiffail="1" mark="3.0000000" >
      <testcode>
                <text>// Testet ob ein for-Konstrukt verwendet wurde</text>
      </testcode>
      <stdin>
                <text></text>
      </stdin>
      <expected>
                <text>Passt!
</text>
      </expected>
      <extra>
                <text><![CDATA[S_A = remove_comments(S_A)
print_if(not (re.search("do", S_A) and re.search("while\\s*\(", S_A)) and re.search("for\\s*\(", S_A), negative_message="Ihr dürft keine DO-WHILE Schleife benutzen, sondern musst ausschließlich ein FOR-Konstrukt schreiben")]]></text>
      </extra>
      <display>
                <text>HIDE_IF_SUCCEED</text>
      </display>
    </testcase>
      <testcase testtype="0" useasexample="1" hiderestiffail="0" mark="1.0000000" >
      <testcode>
                <text>main();</text>
      </testcode>
      <stdin>
                <text>3</text>
      </stdin>
      <expected>
                <text>Wieviele Sekunden soll ich warten?
time 3
time 2
time 1
time 0
</text>
      </expected>
      <extra>
                <text><![CDATA[S_A = remove_comments(S_A)
print_if(list_of_globals(S_A) == [], "", "Nicht erlaubt: Folgende globale Variablen wurden definiert:\n" + "\n".join(list_of_globals(S_A)) + "\n\n")]]></text>
      </extra>
      <display>
                <text>SHOW</text>
      </display>
    </testcase>
      <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000" >
      <testcode>
                <text>main();</text>
      </testcode>
      <stdin>
                <text>0</text>
      </stdin>
      <expected>
                <text>Wieviele Sekunden soll ich warten?
time 0
</text>
      </expected>
      <extra>
                <text><![CDATA[S_A = remove_comments(S_A)
print_if(list_of_globals(S_A) == [], "", "Nicht erlaubt: Folgende globale Variablen wurden definiert:\n" + "\n".join(list_of_globals(S_A)) + "\n\n")]]></text>
      </extra>
      <display>
                <text>HIDE_IF_SUCCEED</text>
      </display>
    </testcase>
      <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000" >
      <testcode>
                <text>main();</text>
      </testcode>
      <stdin>
                <text>-10</text>
      </stdin>
      <expected>
                <text>Wieviele Sekunden soll ich warten?
</text>
      </expected>
      <extra>
                <text><![CDATA[S_A = remove_comments(S_A)
print_if(list_of_globals(S_A) == [], "", "Nicht erlaubt: Folgende globale Variablen wurden definiert:\n" + "\n".join(list_of_globals(S_A)) + "\n\n")]]></text>
      </extra>
      <display>
                <text>HIDE_IF_SUCCEED</text>
      </display>
    </testcase>
    </testcases>
  </question>

</quiz>