<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 2174109  -->
  <question type="coderunner">
    <name>
      <text>Nächste Bitte</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>Schreibt ein Programm, das eine Art Warteschlangenabarbeitung implementiert. Die Warteschlange wird anhand eines Arrays von Zahlen modelliert. Jede Zahl in diesem Array stehe für eine Wartezeit.<br><br></p><ul><li>Legt ein Array mit 5 Elementen in <b>main</b> an und lasst es vom Programmnutzer mit zufälligen Zahlen per Eingabe füllen. </li><li>Schreibt eine Funktion <b>findIndexOfMax</b>, die ein Array mit Zahlen übergeben bekommt und die den Index des Elementes mit der grössten Zahl zurückgibt.</li><li>Aus <b>main</b> soll diese Funktion in einer Schleife aufgerufen werden und dann jeweils das Element des zurückgegebenen Indexes auf <b>INT_MIN</b> gesetzt werden. Ein Element, das als Wert <b>INT_MIN</b> enthält, gelte als abgearbeitet.</li><li><b>findIndexOfMax</b> muss, falls die Warteschlange leer ist (alle Einträge entsprechen <b>INT_MIN</b>), einen definierten Wert zurückgeben (z.B.: <b>UINT_MAX</b>)</li><li>Wird dieser definierte Wert an <b>main</b> zurückgegeben, ist die Warteschlange abgearbeitet und das Programm wird beendet.</li></ul><p>Was ist <b>INT_MIN</b> und <b>UINT_MAX</b>? Schaut hier nach: <a href="https://www.tutorialspoint.com/c_standard_library/limits_h.htm">https://www.tutorialspoint.com/c_standard_library/limits_h.htm</a>.<br><br></p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text></text>
    </generalfeedback>
    <defaultgrade>4</defaultgrade>
    <penalty>0</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <coderunnertype>c_program</coderunnertype>
    <prototypetype>0</prototypetype>
    <allornothing>1</allornothing>
    <penaltyregime>0, 10, 20, ...</penaltyregime>
    <precheck>0</precheck>
    <hidecheck>0</hidecheck>
    <showsource>0</showsource>
    <answerboxlines>35</answerboxlines>
    <answerboxcolumns>100</answerboxcolumns>
    <answerpreload></answerpreload>
    <globalextra></globalextra>
    <useace></useace>
    <resultcolumns></resultcolumns>
    <template></template>
    <iscombinatortemplate></iscombinatortemplate>
    <allowmultiplestdins></allowmultiplestdins>
    <answer><![CDATA[#include <stdio.h>
#include <limits.h>

unsigned int findIndexOfMax(int numbers[], int count)
{
    int max = INT_MIN;
    unsigned int index = UINT_MAX;
    for(unsigned int i = 0; i < count; i++)
    {
        if (max < numbers[i]) 
        { 
            max = numbers[i];
            index = i;
        }
    }
    return index;
}

int main()
{
    unsigned int size = 5;
    int numbers[size];
    printf("Bitte geben Sie %u ganze Zahlen ein: ", size);
    for(unsigned int i = 0; i < size; i++) scanf("%i", &numbers[i]);
    for(unsigned int maxIdx = findIndexOfMax(numbers, size); maxIdx != UINT_MAX; maxIdx = findIndexOfMax(numbers, size))
    {
        printf("Naechstgroesstes: %i \n", numbers[maxIdx]);
        numbers[maxIdx] = INT_MIN;
    }
}]]></answer>
    <validateonsave>1</validateonsave>
    <testsplitterre></testsplitterre>
    <language></language>
    <acelang></acelang>
    <sandbox></sandbox>
    <grader>NearEqualityGrader</grader>
    <cputimelimitsecs></cputimelimitsecs>
    <memlimitmb></memlimitmb>
    <sandboxparams></sandboxparams>
    <templateparams></templateparams>
    <hoisttemplateparams>1</hoisttemplateparams>
    <templateparamslang>twig</templateparamslang>
    <templateparamsevalpertry>0</templateparamsevalpertry>
    <templateparamsevald></templateparamsevald>
    <twigall>0</twigall>
    <uiplugin></uiplugin>
    <uiparameters></uiparameters>
    <attachments>0</attachments>
    <attachmentsrequired>0</attachmentsrequired>
    <maxfilesize>10240</maxfilesize>
    <filenamesregex></filenamesregex>
    <filenamesexplain></filenamesexplain>
    <displayfeedback>1</displayfeedback>
    <giveupallowed>0</giveupallowed>
    <prototypeextra></prototypeextra>
    <testcases>
      <testcase testtype="0" useasexample="1" hiderestiffail="0" mark="1.0000000" >
      <testcode>
                <text>Test 1</text>
      </testcode>
      <stdin>
                <text>1 -2 3 -4 1234</text>
      </stdin>
      <expected>
                <text>Bitte geben Sie 5 ganze Zahlen ein: Naechstgroesstes: 1234
Naechstgroesstes: 3
Naechstgroesstes: 1
Naechstgroesstes: -2
Naechstgroesstes: -4
</text>
      </expected>
      <extra>
                <text></text>
      </extra>
      <display>
                <text>SHOW</text>
      </display>
    </testcase>
      <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000" >
      <testcode>
                <text>Test 2</text>
      </testcode>
      <stdin>
                <text>10 5 20 -100 0</text>
      </stdin>
      <expected>
                <text>Bitte geben Sie 5 ganze Zahlen ein: Naechstgroesstes: 20
Naechstgroesstes: 10
Naechstgroesstes: 5
Naechstgroesstes: 0
Naechstgroesstes: -100
</text>
      </expected>
      <extra>
                <text></text>
      </extra>
      <display>
                <text>HIDE_IF_SUCCEED</text>
      </display>
    </testcase>
      <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000" >
      <testcode>
                <text>Test 3</text>
      </testcode>
      <stdin>
                <text>1 2 3 4 5</text>
      </stdin>
      <expected>
                <text>Bitte geben Sie 5 ganze Zahlen ein: Naechstgroesstes: 5
Naechstgroesstes: 4
Naechstgroesstes: 3
Naechstgroesstes: 2
Naechstgroesstes: 1
</text>
      </expected>
      <extra>
                <text></text>
      </extra>
      <display>
                <text>HIDE_IF_SUCCEED</text>
      </display>
    </testcase>
      <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000" >
      <testcode>
                <text>Test 4</text>
      </testcode>
      <stdin>
                <text>-1 -2 -3 -4 -5</text>
      </stdin>
      <expected>
                <text>Bitte geben Sie 5 ganze Zahlen ein: Naechstgroesstes: -1
Naechstgroesstes: -2
Naechstgroesstes: -3
Naechstgroesstes: -4
Naechstgroesstes: -5
</text>
      </expected>
      <extra>
                <text></text>
      </extra>
      <display>
                <text>HIDE_IF_SUCCEED</text>
      </display>
    </testcase>
    </testcases>
  </question>

</quiz>