<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 3419611  -->
  <question type="coderunner">
    <name>
      <text>Ersetze if/switch (Neu; FERTIG)</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>Das folgende Programm gibt dem/der BenutzerIn verschiedene Optionen zur Auswahl -- also eine Art Menü, aus dem man einzelne Funktionen aufruft. <br><br>Das Programm funktioniert so wie gegeben, der Code für die Auswahl ginge aber auch ein klein bisschen <strong>kompakter </strong>mit einem <strong>switch</strong>-Konstrukt. Ersetzt das <strong>if</strong>-Konstrukt entsprechend. Die Funktionalität soll sich dadurch nicht verändern.<br></p><p><span class="" style="color: rgb(197, 14, 31);">Hinweis</span><span class="" style="color: rgb(197, 14, 31);">:</span> Ihr könnt, wenn ihr wollt, bei dieser Aufgabe eine main-Funktion in eurer Lösung haben. Beim Testen wird jedoch eine (für euch) versteckte <strong>main erstellt</strong>, die eure Funktion menue() aufruft und testet.<br><br></p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <coderunnertype>c_advanced</coderunnertype>
    <prototypetype>0</prototypetype>
    <allornothing>0</allornothing>
    <penaltyregime>0, 0, 5, 10, 20, ...</penaltyregime>
    <precheck>0</precheck>
    <hidecheck>0</hidecheck>
    <showsource>0</showsource>
    <answerboxlines>18</answerboxlines>
    <answerboxcolumns>100</answerboxcolumns>
    <answerpreload><![CDATA[#include <stdio.h>

void menue()
{
    printf("Bitte auswaehlen...\n"
           "\tPPR ist [t]oll\n"
           "\tPPR [s]tinkt mir\n"
           "\tPPR, Flasche [l]eer.\n\n");
    char eingabe = getchar();
    getchar(); // "enter" aus dem Tastaturpuffer entfernen
    if(eingabe == 't')
    {
        printf("Wir moegen dich auch.\n");
    }
    else if(eingabe == 's')
    {
        printf("Hauptsache wir stinken geruchstechnisch nicht.\n");
    }
    else if(eingabe == 'l')
    {
        printf("Giovanni, lernst du Programmieren?\n");
    }
    else
    {
        printf("Das stand gar nicht zur Auswahl. Auf wiedersehen.\n");
    }
}
]]></answerpreload>
    <globalextra></globalextra>
    <useace></useace>
    <resultcolumns></resultcolumns>
    <template></template>
    <iscombinatortemplate></iscombinatortemplate>
    <allowmultiplestdins></allowmultiplestdins>
    <answer><![CDATA[#include <stdio.h>

void menue()
{
    printf("Bitte auswaehlen...\n"
           "\tPPR ist [t]oll\n"
           "\tPPR [s]tinkt mir\n"
           "\tPPR, Flasche [l]eer.\n\n");
    char eingabe = getchar();
    getchar(); // "enter" aus dem Tastaturpuffer entfernen
    switch (eingabe)
    {
        case 't':
            printf("Wir moegen dich auch.\n");
            break;
        case 's':
            printf("Hauptsache wir stinken geruchstechnisch nicht.\n");
            break;
        case 'l':
            printf("Giovanni, lernst du Programmieren?\n");
            break;
        default:
            printf("Das stand gar nicht zur Auswahl. Auf wiedersehen.\n");
            eingabe = 0;
    }
}]]></answer>
    <validateonsave>1</validateonsave>
    <testsplitterre></testsplitterre>
    <language></language>
    <acelang></acelang>
    <sandbox></sandbox>
    <grader></grader>
    <cputimelimitsecs></cputimelimitsecs>
    <memlimitmb></memlimitmb>
    <sandboxparams></sandboxparams>
    <templateparams></templateparams>
    <hoisttemplateparams>1</hoisttemplateparams>
    <templateparamslang>twig</templateparamslang>
    <templateparamsevalpertry>0</templateparamsevalpertry>
    <templateparamsevald>{}</templateparamsevald>
    <twigall>0</twigall>
    <uiplugin></uiplugin>
    <uiparameters></uiparameters>
    <attachments>0</attachments>
    <attachmentsrequired>0</attachmentsrequired>
    <maxfilesize>10240</maxfilesize>
    <filenamesregex></filenamesregex>
    <filenamesexplain></filenamesexplain>
    <displayfeedback>1</displayfeedback>
    <giveupallowed>0</giveupallowed>
    <prototypeextra></prototypeextra>
    <testcases>
      <testcase testtype="0" useasexample="0" hiderestiffail="1" mark="4.0000000" >
      <testcode>
                <text>// Testet ob ein switch-Konstrukt verwendet wurde</text>
      </testcode>
      <stdin>
                <text></text>
      </stdin>
      <expected>
                <text>Passt!
</text>
      </expected>
      <extra>
                <text><![CDATA[S_A = remove_comments(S_A)
print_if(not re.search("\\s+if\\s*", S_A) and re.search("\\s+switch\\s*\(", S_A), negative_message="Ihr dürft keine IF-Verzweigungen benutzen, sondern musst ausschließlich ein SWITCH-Konstrukt schreiben")]]></text>
      </extra>
      <display>
                <text>HIDE_IF_SUCCEED</text>
      </display>
    </testcase>
      <testcase testtype="0" useasexample="1" hiderestiffail="0" mark="1.0000000" >
      <testcode>
                <text>menue();</text>
      </testcode>
      <stdin>
                <text>t</text>
      </stdin>
      <expected>
                <text>Bitte auswaehlen...
\tPPR ist [t]oll
\tPPR [s]tinkt mir
\tPPR, Flasche [l]eer.

Wir moegen dich auch.
</text>
      </expected>
      <extra>
                <text><![CDATA[print_if(list_of_globals(S_A) == [], "", "Nicht erlaubt: Folgende globale Variablen wurden definiert:\n" + "\n".join(list_of_globals(S_A)) + "\n\n")]]></text>
      </extra>
      <display>
                <text>SHOW</text>
      </display>
    </testcase>
      <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000" >
      <testcode>
                <text>menue();</text>
      </testcode>
      <stdin>
                <text>s</text>
      </stdin>
      <expected>
                <text>Bitte auswaehlen...
\tPPR ist [t]oll
\tPPR [s]tinkt mir
\tPPR, Flasche [l]eer.

Hauptsache wir stinken geruchstechnisch nicht.
</text>
      </expected>
      <extra>
                <text><![CDATA[print_if(list_of_globals(S_A) == [], "", "Nicht erlaubt: Folgende globale Variablen wurden definiert:\n" + "\n".join(list_of_globals(S_A)) + "\n\n")]]></text>
      </extra>
      <display>
                <text>HIDE_IF_SUCCEED</text>
      </display>
    </testcase>
      <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000" >
      <testcode>
                <text>menue();</text>
      </testcode>
      <stdin>
                <text>l</text>
      </stdin>
      <expected>
                <text>Bitte auswaehlen...
\tPPR ist [t]oll
\tPPR [s]tinkt mir
\tPPR, Flasche [l]eer.

Giovanni, lernst du Programmieren?
</text>
      </expected>
      <extra>
                <text><![CDATA[print_if(list_of_globals(S_A) == [], "", "Nicht erlaubt: Folgende globale Variablen wurden definiert:\n" + "\n".join(list_of_globals(S_A)) + "\n\n")]]></text>
      </extra>
      <display>
                <text>HIDE_IF_SUCCEED</text>
      </display>
    </testcase>
      <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000" >
      <testcode>
                <text>menue();</text>
      </testcode>
      <stdin>
                <text>a</text>
      </stdin>
      <expected>
                <text>Bitte auswaehlen...
\tPPR ist [t]oll
\tPPR [s]tinkt mir
\tPPR, Flasche [l]eer.

Das stand gar nicht zur Auswahl. Auf wiedersehen.
</text>
      </expected>
      <extra>
                <text><![CDATA[print_if(list_of_globals(S_A) == [], "", "Nicht erlaubt: Folgende globale Variablen wurden definiert:\n" + "\n".join(list_of_globals(S_A)) + "\n\n")]]></text>
      </extra>
      <display>
                <text>HIDE_IF_SUCCEED</text>
      </display>
    </testcase>
    </testcases>
  </question>

</quiz>